[1488715149] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 334.076ms, lost 0%
[1488715209] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 217.593ms, lost 0%
[1488715269] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 300.020ms, lost 0%
[1488715709] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 90.074ms, lost 20%
[1488715759] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;OK;SOFT;2;OK - 187.60.120.9: rta 84.408ms, lost 0%
[1488717669] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 185.394ms, lost 0%
[1488717969] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 345.296ms, lost 0%
[1488718029] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 351.203ms, lost 0%
[1488718089] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 273.582ms, lost 0%
[1488721214] SERVICE ALERT: disponibilidade_nwavgma_net;Check_rot_net_181.222.26.192;WARNING;SOFT;1;WARNING - 181.222.26.192: rta 512.088ms, lost 33%
[1488721264] SERVICE ALERT: disponibilidade_nwavgma_net;Check_rot_net_181.222.26.192;OK;SOFT;2;OK - 181.222.26.192: rta 267.970ms, lost 0%
[1488721389] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 195.943ms, lost 0%
[1488721689] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 296.820ms, lost 0%
[1488721749] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 221.773ms, lost 0%
[1488721809] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 281.835ms, lost 0%
[1488724769] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 91.151ms, lost 20%
[1488724819] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;OK;SOFT;2;OK - 187.60.120.9: rta 91.194ms, lost 0%
[1488724874] SERVICE ALERT: disponibilidade_nwavgma_net;Check_rot_net_181.222.26.192;CRITICAL;SOFT;1;CRITICAL - 181.222.26.192: rta 7438.317ms, lost 0%
[1488724924] SERVICE ALERT: disponibilidade_nwavgma_net;Check_rot_net_181.222.26.192;OK;SOFT;2;OK - 181.222.26.192: rta 334.423ms, lost 0%
[1488725054] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 92.205ms, lost 33%
[1488725104] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;OK;SOFT;2;OK - 187.60.120.9: rta 89.486ms, lost 0%
[1488726329] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 86.090ms, lost 20%
[1488726379] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;OK;SOFT;2;OK - 187.60.120.9: rta 86.074ms, lost 0%
[1488727214] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 92.061ms, lost 33%
[1488727274] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;WARNING;SOFT;2;WARNING - 187.60.120.9: rta 89.293ms, lost 33%
[1488727324] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;OK;SOFT;3;OK - 187.60.120.9: rta 92.078ms, lost 0%
[1488728109] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 191.216ms, lost 0%
[1488728409] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 289.660ms, lost 0%
[1488728469] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 240.836ms, lost 0%
[1488728529] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 352.401ms, lost 0%
[1488728789] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 86.812ms, lost 20%
[1488728849] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;WARNING;SOFT;2;WARNING - 187.60.120.9: rta 87.216ms, lost 20%
[1488728899] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;OK;SOFT;3;OK - 187.60.120.9: rta 84.738ms, lost 0%
[1488731534] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 88.718ms, lost 33%
[1488731584] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;OK;SOFT;2;OK - 187.60.120.9: rta 84.038ms, lost 0%
[1488731609] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 86.090ms, lost 20%
[1488731659] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;OK;SOFT;2;OK - 187.60.120.9: rta 85.695ms, lost 0%
[1488731829] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 192.844ms, lost 0%
[1488732129] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 223.736ms, lost 0%
[1488732189] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 353.573ms, lost 0%
[1488732249] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 242.890ms, lost 0%
[1488732849] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 161.921ms, lost 0%
[1488733749] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 256.347ms, lost 0%
[1488733809] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 311.996ms, lost 0%
[1488733869] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 326.225ms, lost 0%
[1488740469] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 175.158ms, lost 0%
[1488740769] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 244.088ms, lost 0%
[1488740829] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 215.541ms, lost 0%
[1488740889] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 339.456ms, lost 0%
[1488741489] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 155.244ms, lost 0%
[1488741789] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 361.546ms, lost 0%
[1488741849] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 342.437ms, lost 0%
[1488741909] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 226.727ms, lost 0%
[1488742211] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 192.889ms, lost 0%
[1488742509] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 295.973ms, lost 0%
[1488742569] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 270.442ms, lost 0%
[1488742629] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 311.391ms, lost 0%
[1488743894] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 91.531ms, lost 33%
[1488743944] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;OK;SOFT;2;OK - 187.60.120.9: rta 85.972ms, lost 0%
[1488744729] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 160.731ms, lost 0%
[1488745029] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 304.095ms, lost 0%
[1488745089] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 333.937ms, lost 0%
[1488745149] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;SOFT;3;OK - 10.33.44.1: rta 179.260ms, lost 0%
[1488745449] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 262.987ms, lost 0%
[1488745509] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 342.560ms, lost 0%
[1488745569] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;SOFT;3;OK - 10.33.44.1: rta 198.242ms, lost 0%
[1488745869] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 316.719ms, lost 0%
[1488745929] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 347.349ms, lost 0%
[1488745989] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 263.555ms, lost 0%
[1488746954] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 85.543ms, lost 33%
[1488747004] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;OK;SOFT;2;OK - 187.60.120.9: rta 83.912ms, lost 0%
[1488749669] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 141.957ms, lost 20%
[1488749719] SERVICE ALERT: disponibilidade_udi_tvn;Disponibilidade;OK;SOFT;2;OK - 187.60.120.9: rta 87.727ms, lost 0%
[1488750789] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 199.245ms, lost 0%
[1488751089] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 342.252ms, lost 0%
[1488751149] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 314.595ms, lost 0%
[1488751209] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 207.569ms, lost 0%
[1488752710] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 198.169ms, lost 0%
[1488753009] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 353.550ms, lost 0%
[1488753070] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 222.769ms, lost 0%
[1488753129] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 266.959ms, lost 0%
[1488753729] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 166.610ms, lost 0%
[1488754329] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 310.265ms, lost 0%
[1488754389] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 353.348ms, lost 0%
[1488754449] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 225.734ms, lost 0%
[1488756914] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 88.634ms, lost 33%
[1488756964] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;OK;SOFT;2;OK - 187.60.120.9: rta 83.668ms, lost 0%
[1488758174] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 88.231ms, lost 33%
[1488758224] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;OK;SOFT;2;OK - 187.60.120.9: rta 84.025ms, lost 0%
[1488759849] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 194.932ms, lost 0%
[1488760149] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 238.549ms, lost 0%
[1488760209] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 274.037ms, lost 0%
[1488760269] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 312.224ms, lost 0%
[1488762069] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 195.492ms, lost 0%
[1488762369] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 252.222ms, lost 0%
[1488762429] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 303.255ms, lost 0%
[1488762489] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 337.893ms, lost 0%
[1488763089] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 198.915ms, lost 0%
[1488763689] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 335.249ms, lost 0%
[1488763749] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;SOFT;2;OK - 10.33.44.1: rta 199.359ms, lost 0%
[1488764049] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 238.648ms, lost 0%
[1488764109] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 259.917ms, lost 0%
[1488764169] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 300.825ms, lost 0%
[1488764234] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;WARNING;SOFT;1;WARNING - 187.60.120.9: rta 85.099ms, lost 33%
[1488764284] SERVICE ALERT: disponibilidade_udi_tvn;Check_rot_187.60.120.9;OK;SOFT;2;OK - 187.60.120.9: rta 86.202ms, lost 0%
[1488765369] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 156.209ms, lost 0%
[1488765669] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 340.709ms, lost 0%
[1488765729] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 217.732ms, lost 0%
[1488765789] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 268.761ms, lost 0%