[1488448755] SERVICE ALERT: disponibilidade_rebocadores_eqt;Check_fw_187.72.195.143;WARNING;SOFT;1;WARNING - 187.72.195.143: rta 154.911ms, lost 33%
[1488448810] SERVICE ALERT: disponibilidade_rebocadores_eqt;Check_fw_187.72.195.143;OK;SOFT;2;OK - 187.72.195.143: rta 155.094ms, lost 0%
[1488452625] SERVICE ALERT: disponibilidade_rebocadores_eqt;Disponibilidade;WARNING;SOFT;1;WARNING - 187.72.195.143: rta 155.173ms, lost 20%
[1488452680] SERVICE ALERT: disponibilidade_rebocadores_eqt;Disponibilidade;OK;SOFT;2;OK - 187.72.195.143: rta 155.105ms, lost 0%
[1488461715] SERVICE ALERT: disponibilidade_rebocadores_eqt;Check_fw_187.72.195.143;WARNING;SOFT;1;WARNING - 187.72.195.143: rta 156.398ms, lost 33%
[1488461770] SERVICE ALERT: disponibilidade_rebocadores_eqt;Check_fw_187.72.195.143;OK;SOFT;2;OK - 187.72.195.143: rta 155.138ms, lost 0%
[1488469400] SERVICE ALERT: disponibilidade_fundacao_emb;Disponibilidade;WARNING;SOFT;1;WARNING - 201.65.76.130: rta 213.127ms, lost 20%
[1488469455] SERVICE ALERT: disponibilidade_fundacao_emb;Disponibilidade;OK;SOFT;2;OK - 201.65.76.130: rta 170.923ms, lost 0%
[1488472075] SERVICE ALERT: disponibilidade_fundacao_rnp;Disponibilidade;WARNING;SOFT;1;WARNING - 200.137.134.162: rta 114.368ms, lost 20%
[1488472130] SERVICE ALERT: disponibilidade_fundacao_rnp;Disponibilidade;OK;SOFT;2;OK - 200.137.134.162: rta 114.154ms, lost 0%
[1488477985] SERVICE ALERT: fw-prod-nelsonw-01;Network-bge0-Wikitelecom;CRITICAL;SOFT;1;(null)
[1488478045] SERVICE ALERT: fw-prod-nelsonw-01;Network-bge0-Wikitelecom;CRITICAL;SOFT;2;(null)
[1488478065] SERVICE ALERT: fw-prod-nelsonw-01;Network-bge1-Embrael;CRITICAL;SOFT;1;(null)
[1488478105] SERVICE ALERT: fw-prod-nelsonw-01;Network-bge0-Wikitelecom;CRITICAL;HARD;3;(null)
[1488478120] SERVICE ALERT: fw-prod-nelsonw-01;Data_e_Hora;WARNING;SOFT;1;CRITICAL - nao consegui capturar valores
[1488478125] SERVICE ALERT: fw-prod-nelsonw-01;Network-bge1-Embrael;CRITICAL;SOFT;2;(null)
[1488478145] SERVICE ALERT: fw-prod-nelsonw-01;Network-em0-Local;CRITICAL;SOFT;1;(null)
[1488478180] SERVICE ALERT: fw-prod-nelsonw-01;Data_e_Hora;WARNING;SOFT;2;CRITICAL - nao consegui capturar valores
[1488478185] SERVICE ALERT: fw-prod-nelsonw-01;Network-bge1-Embrael;CRITICAL;HARD;3;(null)
[1488478190] SERVICE ALERT: fw-prod-nelsonw-01;CPU;CRITICAL;SOFT;1;(Service Check Timed Out)
[1488478205] SERVICE ALERT: fw-prod-nelsonw-01;Network-em0-Local;CRITICAL;SOFT;2;(null)
[1488478240] SERVICE ALERT: fw-prod-nelsonw-01;Data_e_Hora;WARNING;HARD;3;CRITICAL - nao consegui capturar valores
[1488478265] SERVICE ALERT: fw-prod-nelsonw-01;Network-em0-Local;CRITICAL;HARD;3;(null)
[1488478270] SERVICE ALERT: fw-prod-nelsonw-01;DISK;CRITICAL;SOFT;1;(Service Check Timed Out)
[1488478365] SERVICE ALERT: fw-prod-nelsonw-01;LOAD;CRITICAL;SOFT;1;(Service Check Timed Out)
[1488478370] SERVICE ALERT: fw-prod-nelsonw-01;LOAD;OK;SOFT;2;LOAD AVERAGE OK - 0.32,0.52,0.48
[1488478375] SERVICE ALERT: fw-prod-nelsonw-01;DISK;OK;SOFT;2;DISK USAGE OK - Usage: 3% (Total: 54.22 GiB - Free: 48.22 GiB)
[1488478395] SERVICE ALERT: fw-prod-nelsonw-01;CPU;OK;SOFT;2;CPU OK - Usage: 4%
[1488478535] SERVICE ALERT: fw-prod-nelsonw-01;Data_e_Hora;OK;HARD;3;Hora atual do Host: 2017-3-2 15:15:3.9 -3:0
[1488478555] SERVICE ALERT: fw-prod-nelsonw-01;Network-em0-Local;OK;HARD;3;Average IN: 0KBits/s (0.00%), Average OUT: 0KBits/s (0.00%)<br>Total RX: 2230.05 MBytes, Total TX: 3129.27 MBytes
[1488479265] SERVICE ALERT: fw-prod-nelsonw-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;1;WARNING - Average IN: 78.25MBits/s (626.01%), Average OUT: 1.79MBits/s (14.30%)<br>Total RX: 3780.00 MBytes, Total TX: 1025.66 MBytes
[1488479325] SERVICE ALERT: fw-prod-nelsonw-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;2;WARNING - Average IN: 78.98MBits/s (631.81%), Average OUT: 1.40MBits/s (11.23%)<br>Total RX: 4253.86 MBytes, Total TX: 1034.08 MBytes
[1488479385] SERVICE ALERT: fw-prod-nelsonw-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;HARD;3;WARNING - Average IN: 15.93MBits/s (127.45%), Average OUT: 1.07MBits/s (8.53%)<br>Total RX: 253.44 MBytes, Total TX: 1040.47 MBytes
[1488479990] SERVICE ALERT: disponibilidade_rebocadores_eqt;Disponibilidade;CRITICAL;SOFT;1;CRITICAL - 187.72.195.143: rta nan, lost 100%
[1488480041] SERVICE ALERT: disponibilidade_rebocadores_eqt;Disponibilidade;OK;SOFT;2;OK - 187.72.195.143: rta 154.959ms, lost 0%
[1488480096] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;1;WARNING - Average IN: 70.38MBits/s (563.04%), Average OUT: 1.43MBits/s (11.46%)<br>Total RX: 3623.89 MBytes, Total TX: 1128.34 MBytes
[1488480156] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;2;WARNING - Average IN: 69.34MBits/s (554.75%), Average OUT: 1.55MBits/s (12.43%)<br>Total RX: 4039.95 MBytes, Total TX: 1137.66 MBytes
[1488480216] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;HARD;3;WARNING - Average IN: 24.86MBits/s (198.85%), Average OUT: 653.7KBits/s (5.11%)<br>Total RX: 4189.09 MBytes, Total TX: 1141.49 MBytes
[1488482916] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;OK;HARD;3;Average IN: 7.47MBits/s (59.79%), Average OUT: 504.8KBits/s (3.94%)<br>Total RX: 3375.25 MBytes, Total TX: 1516.13 MBytes
[1488482951] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 325.743ms, lost 20%
[1488483011] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 273.654ms, lost 0%
[1488483071] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;SOFT;3;OK - 10.33.44.1: rta 152.719ms, lost 0%
[1488483216] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;1;WARNING - Average IN: 24.97MBits/s (199.73%), Average OUT: 1.03MBits/s (8.22%)<br>Total RX: 4129.24 MBytes, Total TX: 1547.15 MBytes
[1488483276] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;2;WARNING - Average IN: 45.45MBits/s (363.63%), Average OUT: 754.1KBits/s (5.89%)<br>Total RX: 301.42 MBytes, Total TX: 1551.49 MBytes
[1488483336] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;OK;SOFT;3;Average IN: 3.80MBits/s (30.36%), Average OUT: 210.5KBits/s (1.64%)<br>Total RX: 324.19 MBytes, Total TX: 1552.72 MBytes
[1488483371] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;1;WARNING - 10.33.44.1: rta 213.248ms, lost 0%
[1488483431] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;SOFT;2;WARNING - 10.33.44.1: rta 215.828ms, lost 0%
[1488483491] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;WARNING;HARD;3;WARNING - 10.33.44.1: rta 215.045ms, lost 0%
[1488483676] SERVICE ALERT: disponibilidade_fundacao_rnp;Check_fw_200.137.134.162;WARNING;SOFT;1;WARNING - 200.137.134.162: rta 171.007ms, lost 33%
[1488483726] SERVICE ALERT: disponibilidade_fundacao_rnp;Check_fw_200.137.134.162;OK;SOFT;2;OK - 200.137.134.162: rta 170.356ms, lost 0%
[1488484836] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;1;WARNING - Average IN: 21.10MBits/s (168.78%), Average OUT: 702KBits/s (5.48%)<br>Total RX: 1546.37 MBytes, Total TX: 1620.34 MBytes
[1488484896] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;2;WARNING - Average IN: 29.90MBits/s (239.19%), Average OUT: 545.7KBits/s (4.26%)<br>Total RX: 1725.76 MBytes, Total TX: 1623.54 MBytes
[1488484956] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;OK;SOFT;3;Average IN: 6.19MBits/s (49.49%), Average OUT: 346.9KBits/s (2.71%)<br>Total RX: 1762.88 MBytes, Total TX: 1625.57 MBytes
[1488485936] SERVICE ALERT: disponibilidade_fundacao_rnp;Disponibilidade;WARNING;SOFT;1;WARNING - 200.137.134.162: rta 170.114ms, lost 20%
[1488485986] SERVICE ALERT: disponibilidade_fundacao_rnp;Disponibilidade;OK;SOFT;2;OK - 200.137.134.162: rta 170.039ms, lost 0%
[1488486756] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;1;WARNING - Average IN: 30.93MBits/s (247.47%), Average OUT: 983KBits/s (7.68%)<br>Total RX: 3514.59 MBytes, Total TX: 1727.27 MBytes
[1488486816] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;2;WARNING - Average IN: 54.69MBits/s (437.51%), Average OUT: 1.41MBits/s (11.28%)<br>Total RX: 3842.72 MBytes, Total TX: 1735.74 MBytes
[1488486876] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;HARD;3;WARNING - Average IN: 40.04MBits/s (320.29%), Average OUT: 1.60MBits/s (12.84%)<br>Total RX: 4082.94 MBytes, Total TX: 1745.37 MBytes
[1488488981] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;OK;HARD;3;Average IN: 5.09MBits/s (40.71%), Average OUT: 1.14MBits/s (9.15%)<br>Total RX: 969.07 MBytes, Total TX: 1996.61 MBytes
[1488491381] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;1;WARNING - Average IN: 48.94MBits/s (391.52%), Average OUT: 1.07MBits/s (8.57%)<br>Total RX: 3773.45 MBytes, Total TX: 2124.05 MBytes
[1488491441] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;2;WARNING - Average IN: 73.85MBits/s (590.77%), Average OUT: 1.67MBits/s (13.39%)<br>Total RX: 4216.53 MBytes, Total TX: 2134.09 MBytes
[1488491501] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;OK;SOFT;3;Average IN: 5.46MBits/s (43.71%), Average OUT: 208.6KBits/s (1.63%)<br>Total RX: 153.31 MBytes, Total TX: 2135.31 MBytes
[1488492816] SERVICE ALERT: disponibilidade_fundacao_emb;Check_fw_201.65.76.130;WARNING;SOFT;1;WARNING - 201.65.76.130: rta 181.460ms, lost 33%
[1488492871] SERVICE ALERT: disponibilidade_fundacao_emb;Check_fw_201.65.76.130;OK;SOFT;2;OK - 201.65.76.130: rta 178.800ms, lost 0%
[1488499901] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;1;WARNING - Average IN: 17.71MBits/s (141.70%), Average OUT: 955.6KBits/s (7.47%)<br>Total RX: 4356.40 MBytes, Total TX: 2400.86 MBytes
[1488499961] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;2;WARNING - Average IN: 39.68MBits/s (317.48%), Average OUT: 2.49MBits/s (19.89%)<br>Total RX: 502.47 MBytes, Total TX: 2416.02 MBytes
[1488500021] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;HARD;3;WARNING - Average IN: 38.55MBits/s (308.38%), Average OUT: 2.66MBits/s (21.31%)<br>Total RX: 737.61 MBytes, Total TX: 2432.27 MBytes
[1488500321] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;OK;HARD;3;Average IN: 7.59MBits/s (60.70%), Average OUT: 694.4KBits/s (5.42%)<br>Total RX: 962.94 MBytes, Total TX: 2452.41 MBytes
[1488502421] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;1;WARNING - Average IN: 39.75MBits/s (317.97%), Average OUT: 1.77MBits/s (14.13%)<br>Total RX: 2648.63 MBytes, Total TX: 2527.53 MBytes
[1488502481] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;2;WARNING - Average IN: 52.18MBits/s (417.47%), Average OUT: 2.07MBits/s (16.55%)<br>Total RX: 2956.51 MBytes, Total TX: 2539.74 MBytes
[1488502541] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;HARD;3;WARNING - Average IN: 38.46MBits/s (307.70%), Average OUT: 2.38MBits/s (19.07%)<br>Total RX: 3187.28 MBytes, Total TX: 2554.04 MBytes
[1488504041] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;OK;HARD;3;Average IN: 2.39MBits/s (19.13%), Average OUT: 127.2KBits/s (0.99%)<br>Total RX: 3790.57 MBytes, Total TX: 2762.94 MBytes
[1488504276] SERVICE ALERT: disponibilidade_fundacao_emb;Check_fw_201.65.76.130;WARNING;SOFT;1;WARNING - 201.65.76.130: rta 180.559ms, lost 33%
[1488504331] SERVICE ALERT: disponibilidade_fundacao_emb;Check_fw_201.65.76.130;OK;SOFT;2;OK - 201.65.76.130: rta 179.390ms, lost 0%
[1488504941] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;1;WARNING - Average IN: 25.14MBits/s (201.12%), Average OUT: 810.2KBits/s (6.33%)<br>Total RX: 4710.18 MBytes, Total TX: 2794.96 MBytes
[1488505001] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;WARNING;SOFT;2;WARNING - Average IN: 20.56MBits/s (164.49%), Average OUT: 1.02MBits/s (8.13%)<br>Total RX: 733.43 MBytes, Total TX: 2800.85 MBytes
[1488505061] SERVICE ALERT: fw-prod-nwavgma-01;Network-em9-net1_192_168_0_2_24_4001;OK;SOFT;3;Average IN: 2.83MBits/s (22.64%), Average OUT: 113.1KBits/s (0.88%)<br>Total RX: 750.42 MBytes, Total TX: 2801.52 MBytes
[1488507191] SERVICE ALERT: fw-prod-nwavgma-01;Disponibilidade;OK;HARD;3;OK - 10.33.44.1: rta 149.818ms, lost 0%